rule_files:
  - ../rules/rhoam_general.yml
tests:
  # The interval is closely tied to the values of the input series.
  # It represents the duration between each value of the metrics.
  - interval: 1m
    # The input series define the values for individual metrics throughout their lifecycle.
    # They should be treated as a whole and not as individual test cases.
    # `_` indicates a missing scrape sample.
    # `stale` indicates a stale sample.
    # `stale _x10` indicates the metric is missing for 10m.
    # `0` is used as a dummy value to imply that the metric is present and its value is `0`.
    # `0 0 0 0 0` series represent a metric that is present for 5m with a value of `0`.
    # To improve readability and address high intervals one can use the expanding notation.
    # `0 0 0 0 0` is equivalent to `0+0x4`.
    # `a+bxc` becomes `a a+b a+(2*b) a+(3*b) … a+(c*b)`.
    # `a-bxc` becomes `a a-b a-(2*b) a-(3*b) … a-(c*b)`.
    input_series:
      # missing for 10m, present for 1m with value of `0`, missing for 10m, present for 85m with value of `0`.
      - series: kube_endpoint_address
        values: "stale _x10 0 stale _x10 0+0x84"
      # missing for 10m, present for 13m with value of `0`, missing for 10m, present for 73m with value of `0`.
      - series: kube_pod_container_status_ready
        values: "stale _x10 0+0x12 stale _x10 0+0x72"
      # missing for 10m, present for 25m with value of `0`, missing for 10m, present for 61m with value of `0`.
      - series: kube_pod_labels
        values: "stale _x10 0+0x24 stale _x10 0+0x60"
      # missing for 10m, present for 37m with value of `0`, missing for 10m, present for 49m with value of `0`.
      - series: kube_pod_status_phase
        values: "stale _x10 0+0x36 stale _x10 0+0x48"
      # missing for 10m, present for 49m with value of `0`, missing for 10m, present for 37m with value of `0`.
      - series: kube_pod_status_ready
        values: "stale _x10 0+0x48 stale _x10 0+0x36"
      # missing for 10m, present for 61m with value of `0`, missing for 10m, present for 25m with value of `0`.
      - series: kube_secret_info
        values: "stale _x10 0+0x60 stale _x10 0+0x24"
      # missing for 10m, present for 73m with value of `0`, missing for 10m, present for 13m with value of `0`.
      - series: rhoam_version
        values: "stale _x10 0+0x72 stale _x10 0+0x12"
      # missing for 10m, present for 85m with value of `0`, missing for 10m, present for 1m with value of `0`.
      - series: threescale_portals
        values: "stale _x10 0+0x84 stale _x10 0"
    # In this test suite we are checking whether metric/s are present or not and fire an alert based on that.
    # If one wants to simulate a metrics
    alert_rule_test:
      # [-] kube_endpoint_address
      # [-] kube_pod_container_status_ready
      # [-] kube_pod_labels
      # [-] kube_pod_status_phase
      # [-] kube_pod_status_ready
      # [-] kube_secret_info
      # [-] rhoam_version
      # [-] threescale_portals
      # alert expected to fire
      - eval_time: 10m
        alertname: CriticalMetricsMissing
        exp_alerts:
          - exp_labels:
              severity: critical
            exp_annotations:
              message: one or more critical metrics have been missing for 10+ minutes

      # [+] kube_endpoint_address
      # [+] kube_pod_container_status_ready
      # [+] kube_pod_labels
      # [+] kube_pod_status_phase
      # [+] kube_pod_status_ready
      # [+] kube_secret_info
      # [+] rhoam_version
      # [+] threescale_portals
      # alert not expected to fire
      - eval_time: 11m
        alertname: CriticalMetricsMissing
        exp_alerts: [ ]

      # [-] kube_endpoint_address
      # [+] kube_pod_container_status_ready
      # [+] kube_pod_labels
      # [+] kube_pod_status_phase
      # [+] kube_pod_status_ready
      # [+] kube_secret_info
      # [+] rhoam_version
      # [+] threescale_portals
      # alert expected to fire
      - eval_time: 22m
        alertname: CriticalMetricsMissing
        exp_alerts:
          - exp_labels:
              severity: critical
            exp_annotations:
              message: one or more critical metrics have been missing for 10+ minutes

      # [+] kube_endpoint_address
      # [-] kube_pod_container_status_ready
      # [+] kube_pod_labels
      # [+] kube_pod_status_phase
      # [+] kube_pod_status_ready
      # [+] kube_secret_info
      # [+] rhoam_version
      # [+] threescale_portals
      # alert not expected to fire
      - eval_time: 34m
        alertname: CriticalMetricsMissing
        exp_alerts:
          - exp_labels:
              severity: critical
            exp_annotations:
              message: one or more critical metrics have been missing for 10+ minutes

      # [+] kube_endpoint_address
      # [+] kube_pod_container_status_ready
      # [-] kube_pod_labels
      # [+] kube_pod_status_phase
      # [+] kube_pod_status_ready
      # [+] kube_secret_info
      # [+] rhoam_version
      # [+] threescale_portals
      # alert not expected to fire
      - eval_time: 46m
        alertname: CriticalMetricsMissing
        exp_alerts:
          - exp_labels:
              severity: critical
            exp_annotations:
              message: one or more critical metrics have been missing for 10+ minutes

      # [+] kube_endpoint_address
      # [+] kube_pod_container_status_ready
      # [+] kube_pod_labels
      # [-] kube_pod_status_phase
      # [+] kube_pod_status_ready
      # [+] kube_secret_info
      # [+] rhoam_version
      # [+] threescale_portals
      # alert not expected to fire
      - eval_time: 58m
        alertname: CriticalMetricsMissing
        exp_alerts:
          - exp_labels:
              severity: critical
            exp_annotations:
              message: one or more critical metrics have been missing for 10+ minutes

      # [+] kube_endpoint_address
      # [+] kube_pod_container_status_ready
      # [+] kube_pod_labels
      # [+] kube_pod_status_phase
      # [-] kube_pod_status_ready
      # [+] kube_secret_info
      # [+] rhoam_version
      # [+] threescale_portals
      # alert not expected to fire
      - eval_time: 70m
        alertname: CriticalMetricsMissing
        exp_alerts:
          - exp_labels:
              severity: critical
            exp_annotations:
              message: one or more critical metrics have been missing for 10+ minutes

      # [+] kube_endpoint_address
      # [+] kube_pod_container_status_ready
      # [+] kube_pod_labels
      # [+] kube_pod_status_phase
      # [+] kube_pod_status_ready
      # [-] kube_secret_info
      # [+] rhoam_version
      # [+] threescale_portals
      # alert not expected to fire
      - eval_time: 82m
        alertname: CriticalMetricsMissing
        exp_alerts:
          - exp_labels:
              severity: critical
            exp_annotations:
              message: one or more critical metrics have been missing for 10+ minutes

      # [+] kube_endpoint_address
      # [+] kube_pod_container_status_ready
      # [+] kube_pod_labels
      # [+] kube_pod_status_phase
      # [+] kube_pod_status_ready
      # [+] kube_secret_info
      # [-] rhoam_version
      # [+] threescale_portals
      # alert not expected to fire
      - eval_time: 94m
        alertname: CriticalMetricsMissing
        exp_alerts:
          - exp_labels:
              severity: critical
            exp_annotations:
              message: one or more critical metrics have been missing for 10+ minutes

      # [+] kube_endpoint_address
      # [+] kube_pod_container_status_ready
      # [+] kube_pod_labels
      # [+] kube_pod_status_phase
      # [+] kube_pod_status_ready
      # [+] kube_secret_info
      # [+] rhoam_version
      # [-] threescale_portals
      # alert not expected to fire
      - eval_time: 106m
        alertname: CriticalMetricsMissing
        exp_alerts:
          - exp_labels:
              severity: critical
            exp_annotations:
              message: one or more critical metrics have been missing for 10+ minutes

      # [+] kube_endpoint_address
      # [+] kube_pod_container_status_ready
      # [+] kube_pod_labels
      # [+] kube_pod_status_phase
      # [+] kube_pod_status_ready
      # [+] kube_secret_info
      # [+] rhoam_version
      # [+] threescale_portals
      # alert not expected to fire
      - eval_time: 107m
        alertname: CriticalMetricsMissing
        exp_alerts: [ ]
